```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

# Homework 06: Testing Hypotheses

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in this course.

**Reading**: 
* [Testing Hypotheses](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

<div style="background-color:#f2d8b2; border-left:6px solid #b36a00; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>ℹ️ Appropriate collaboration and AI usage</b><br>
<span style="font-size:0.9em;">Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.</span>
</div>

<div style="background-color:#dac1d6; border-left:6px solid #853175; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>⚠️ This homework has hidden tests in it!</b><br>
<span style="font-size:0.9em;">Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**</span>
</div>

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Part 1. Vaccinations Across the Nation

A vaccination clinic has two types of vaccines against a disease. Each person who comes in to be vaccinated gets either Vaccine 1 or Vaccine 2. One week, everyone who came in on Monday, Wednesday, and Friday was given Vaccine 1. Everyone who came in on Tuesday and Thursday was given Vaccine 2. The clinic is closed on weekends.

Doctor Gibson at the clinic said, "Oh wow, it's just like tossing a coin that lands heads with chance $\frac{3}{5}$. Heads you get Vaccine 1 and Tails you get Vaccine 2."

But Doctor Stutzman said, "No, it's not. We're not doing anything like tossing a coin."

That week, the clinic gave Vaccine 1 to 211 people and Vaccine 2 to 107 people. Conduct a test of hypotheses to see which doctor's position is better supported by the data.

### Question 1.1

Given the information above, what was the sample size for the data, and what was the percentage of people who got **Vaccine 1?**

*Note*: Your percent should be a number between 0 and 100.

```python
sample_size = ...
percent_V1 = ...

print(f"Sample Size: {sample_size}")
print(f"Vaccine 1 Percent: {percent_V1}")
```

```python
grader.check("q1_1")
```

<!-- BEGIN QUESTION -->

### Question 1.2.

State the null hypothesis. It should reflect the position of either Dr. Gibson or Dr. Stutzman.

*Note:* Check out [11.3](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-1-the-hypotheses) for a refresher on hypotheses.

<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.3.

State the alternative hypothesis. It should reflect the position of the doctor you did not choose to represent in Question 1.2.

*Note:* Check out [11.3](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-1-the-hypotheses) for a refresher on hypotheses.


<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.4.

Doctor Gibson would like to use the absolute difference between the percent of people who got Vaccine 1 and 60% as the statistic to test the hypotheses. List at least one additional reasonable choice for a statistic that could be used to test these hypotheses, and briefly explain why your choice is reasonable.

*Hint*: For a refresher on choosing test statistics, check out the textbook section on [Test Statistics](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html).


<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

### Question 1.5.

Using your answer from Questions 1.1 and 1.4, find the observed value of the test statistic and assign it to the variable `observed_statistic`.

```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q1_5")
```

### Question 1.6.

In order to perform this hypothesis test, you must simulate the test statistic. From the four options below, pick the assumption that is needed for this simulation. Assign `assumption_needed` to an integer corresponding to the assumption.

1. The statistic must be simulated under the null hypothesis.
2. The statistic must be simulated under the alternative hypothesis.
3. The statistic must be simulated under both hypotheses.
4. No assumptions are needed. We can just simulate the statistic.

```python
assumption_needed = ...
assumption_needed
```

```python
grader.check("q1_6")
```

<!-- BEGIN QUESTION -->

### Question 1.7.

Simulate 20,000 values of the test statistic under the assumption you picked in Question 1.6. 

As usual, start by defining a function that simulates one value of the statistic. Your function should use `sample_proportions`. (You may find a variable defined in Question 1.1 useful here!) Then, write a `for` loop to simulate multiple values and collect them in the array `simulated_statistics`.

Use as many lines of code as you need. You will visualize the distribution of these simulated values in Question 1.8.

```python
def one_simulated_statistic():
    ...

num_simulations = 20000

simulated_statistics = ...
for ... in ...:
    ...

# Run this cell a few times to see how the simulated statistic changes
one_simulated_statistic()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.8.

Using the results from your simulation, generate a histogram of the empirical distribution of the simulated test statistics. The red dot marks the observed statistic you found in Question 1.5.


```python
...
plt.scatter(observed_statistic, -0.002, color='red', s=40);
```

<!-- END QUESTION -->

### Question 1.9.

Using `simulated_statistics`, `observed_statistic`, and `num_simulations`, find the empirical p-value based on the simulation.

```python
p_value = ...
p_value
```

```python
grader.check("q1_9")
```

### Question 1.10.

Assign `correct_doctor` to the number corresponding to the correct statement below. Use the 5% cutoff for the p-value.

1. The data support Dr. Gibson's position more than they support Dr. Stutzman's.
2. The data support Dr. Stutzman's position more than they support Dr. Gibson's.

As a reminder, here are the two claims made by Dr. Gibson and Dr. Stutzman:
> **Doctor Gibson:** "Oh wow, it's just like tossing a coin that lands heads with chance $\frac{3}{5}$. Heads you get Vaccine 1 and Tails you get Vaccine 2."

>**Doctor Stutzman:** "No, it's not. We're not doing anything like tossing a coin."


```python
correct_doctor = ...
correct_doctor
```

```python
grader.check("q1_10")
```

## Part 2: Using TVD as a Test Statistic

Before beginning this section, please read [this section](https://inferentialthinking.com/chapters/11/2/Multiple_Categories.html#a-new-statistic-the-distance-between-two-distributions) of the textbook on TVD!

**Total variation distance (TVD)** is a special type of test statistic that we use when we want to compare two distributions of *categorical data*. It is often used when we observe that a set of observed proportions/probabilities is different than what we expect under the null model. 

Consider a six-sided die that we roll 6,000 times. If the die is fair, we would expect that each face comes up $\frac{1}{6}$ of the time. By random chance, a fair die won't always result in equal proportions (that is, we won't get exactly 1,000 of each face). However, if we suspect that the die might be unfair based on the data, we can conduct a hypothesis test using TVD to compare the expected [$\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$] distribution to what is actually observed.

In this part of the homework, we'll look at how we can use TVD to determine the effect that different factors have on happiness. 

We will be working with data from the [Gallup World Poll](https://www.gallup.com/analytics/349487/gallup-global-happiness-center.aspx#:~:text=World%20Happiness%20Report&text=Using%20the%20Gallup%20World%20Poll,about%20the%20World%20Happiness%20Report.) that is presented in the World Happiness Report, a survey of the state of global happiness. The survey ranked 155 countries by overall happiness and estimated the influence that economic production, social support, life expectancy, freedom, absence of corruption, and generosity had on population happiness. The study has been repeated for several years, but we'll be looking at data from the 2016 survey.

Run the cell below to load in the `happiness_scores` table.

```python
happiness_scores = Table.read_table("happiness_scores.csv")
happiness_scores.show(5)
```

Participants in the study were asked to evaluate their life satisfaction on a scale of 0 (worst possible life) to 10 (best possible life). The responses for each country were averaged to create the `Happiness Score`.

The columns `Economy (GDP per Capita)`, `Family`, `Health (Life Expectancy)`, `Freedom`, `Trust (Government Corruption)`, and `Generosity` estimate the extent to which each factor influences happiness, for better or for worse. The happiness score is the sum of these factors; the larger a factor is, the more it contributes to overall happiness. [In other words, if you add up all the factors (in addition to a "Difference from Dystopia" value we excluded in the dataset), you get the happiness score.]

Let's look at the different factors that affect happiness in the United States. Run the cell below to view the row in `us_happiness` that contains data for the United States.

```python
us_happiness = happiness_scores.where("Country", "United States")
us_happiness
```

**To compare the different factors, we'll look at the proportion of the happiness score that is attributed to each variable. 
You can find these proportions in the table `us_happiness_factors` after running the cell below.**

*Note:* The factors shown in `us_happiness` don't add up exactly to the happiness score, so we adjusted the proportions to  only account for the data we have access to. The proportions were found by dividing each Happiness Factor value by the sum of all Happiness Factor values in `us_happiness`.

```python
us_happiness_factors = Table().read_table("us_happiness_factors.csv")
us_happiness_factors
```

<!-- BEGIN QUESTION -->

### Question 2.1

Suppose we want to test whether or not each factor contributes the same amount to the overall Happiness Score. Define the null hypothesis, alternative hypothesis, and test statistic in the cell below. Feel free to check your work with another student or your teacher.

*Note:* Please format your answer as follows:
- Null Hypothesis: ...  
- Alternative Hypothesis: ...  
- Test Statistic: ...  

<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

### Question 2.2

Write a function `calculate_tvd` that takes in the observed distribution (`obs_dist`) and expected distribution under the null hypothesis (`null_dist`) and calculates the total variation distance. Use this function to set `observed_tvd` to be equal to the observed test statistic.

```python
null_distribution = np.ones(6) * (1/6)

def calculate_tvd(obs_dist, null_dist):
    ...
    
observed_tvd = ...
observed_tvd
```

```python
grader.check("q2_2")
```

### Question 2.3
Create an array called `simulated_tvds` that contains 10,000 simulated values under the null hypothesis. Assume that the original sample consisted of 1,000 individuals.

*Hint:* The `sample_proportions` function may be helpful to you. Refer to the [Python Reference Sheet](https://www.data8.org/reference/#:~:text=sample_proportions%28sample_size%2C%20model_proportions%29) to read up on it!

```python
simulated_tvds = ...

...
```

```python
grader.check("q2_3")
```

Run the cell below to plot a histogram of your simulated test statistics, as well as a red dot representing the observed value of the test statistic.

```python
Table().with_column("Simulated TVDs", simulated_tvds).hist()
plt.scatter(observed_tvd, -0.001, color='red', s=70, zorder=2);
plt.show();
```

### Question 2.4.

Use your simulated statistics to calculate the p-value of your test. Make sure that this number is consistent with what you observed in the histogram above.

```python
p_value_tvd = ...
p_value_tvd
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

### Question 2.5.

What can you conclude about how each factor contributes to the overall happiness score in the US? Explain your answer using the results of your hypothesis test. Assume a p-value cutoff of 5%.

<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

## Part 3: Who is Older?

Data scientists have drawn a simple random sample of size 500 from a large population of adults. Each member of the population happened to identify as either "male" or "female". Data was collected on several attributes of the sampled people, including age. The table `sampled_ages` contains one row for each person in the sample, with columns containing the individual's gender identity and age.

```python
sampled_ages = Table.read_table('age.csv')
sampled_ages.show(5)
```

### Question 3.1.
How many females were there in our sample? Please use the provided skeleton code.

*Hint:* Keep in mind that `.group` sorts categories in alphabetical order!


```python
num_females = sampled_ages.group(...)...
num_females
```

```python
grader.check("q3_1")
```

### Question 3.2.

Complete the cell below so that `avg_male_vs_female` evaluates to `True` if the sampled males are older than the sampled females on average, and `False` otherwise. Use Python code to achieve this.

```python
group_mean_tbl = sampled_ages.group(...)
group_means = group_mean_tbl...       # array of mean ages
avg_male_vs_female = group_means... > group_means...
avg_male_vs_female
```

```python
grader.check("q3_2")
```

### Question 3.3.

The data scientists want to use the data to test whether males are older than females—or, in other words, whether the ages of the two groups have the same distribution. One of the following statements is their null hypothesis and another is their alternative hypothesis. Assign `null_statement_number` and `alternative_statement_number` to the numbers of the correct statements in the code cell below.

1. In the sample, the males and females have the same distribution of ages; the sample averages of the two groups are different due to chance.
2. In the population, the males and females have the same distribution of ages; the sample averages of the two groups are different due to chance.
3. The age distributions of males and females in the population are different due to chance.
4. The males in the sample are older than the females, on average.
5. The males in the population are older than the females, on average.
6. The average ages of the males and females in the population are different.

```python
null_statement_number = ...
alternative_statement_number = ...
```

```python
grader.check("q3_3")
```

### Question 3.4.

The data scientists have decided to use a permutation test. Assign `permutation_test_reason` to the number corresponding to the reason they made this choice.

1. Since a person's age shouldn't be related to their gender, it doesn't matter who is labeled "male" and who is labeled "female", so you can use permutations.
2. Under the null hypothesis, permuting the labels in the `sampled_ages` table is equivalent to drawing a new random sample with the same number of males and females as in the original sample.
3. Under the null hypothesis, permuting the rows of `sampled_ages` table is equivalent to drawing a new random sample with the same number of males and females as in the original sample.

```python
permutation_test_reason = ...
permutation_test_reason
```

```python
grader.check("q3_4")
```

### Question 3.5.

To test their hypotheses, the data scientists have followed our textbook's advice and chosen a test statistic where the following statement is true: Large values of the test statistic favor the alternative hypothesis.

The data scientists' test statistic is one of the two options below. Which one is it? Assign the appropriate number to the variable `correct_test_stat`.

1. "male age average - female age average" in a sample created by randomly shuffling the male/female labels
2. "|male age average - female age average|" in a sample created by randomly shuffling the male/female labels

```python
correct_test_stat = ...
correct_test_stat
```

```python
grader.check("q3_5")
```

### Question 3.6.

Complete the cell below so that `observed_statistic_ab` evaluates to the observed value of the data scientists' test statistic. Use as many lines of code as you need, and remember that you can use any quantity, table, or array that you created earlier.

```python
observed_statistic_ab = ...
observed_statistic_ab
```

```python
grader.check("q3_6")
```

### Question 3.7. 

Assign `shuffled_labels` to an array of shuffled male/female labels. The rest of the code puts the array in a table along with the data in `sampled_ages`.

*Note:* Check out [12.1](https://inferentialthinking.com/chapters/12/1/AB_Testing.html#predicting-the-statistic-under-the-null-hypothesis) for a refresher on random permutations.


```python
shuffled_labels = ...
original_with_shuffled_labels = sampled_ages.with_columns('Shuffled Label', shuffled_labels)
original_with_shuffled_labels
```

```python
grader.check("q3_7")
```

### Question 3.8.

The comparison below uses the array `shuffled_labels` from Question 3.7 and the count `num_females` from Question 3.1.

For this comparison, assign the correct number from one of the following options to the variable `correct_q8`. **Pretend this is a midterm problem and solve it without doing the calculation in a code cell.**

`comp = np.count_nonzero(shuffled_labels == 'female') == num_females`

1. `comp` is set to `True`.
2. `comp` is set to `False`.
3. `comp` is set to `True` or `False`, depending on how the shuffle came out.


```python
correct_q8 = ...
correct_q8
```

```python
grader.check("q3_8")
```

### Question 3.9.

Define a function `simulate_one_statistic` that takes no arguments and returns one simulated value of the test statistic. We've given you a skeleton, but feel free to approach this question in a way that makes sense to you. Use as many lines of code as you need. Refer to the code you have previously written in this problem, as you might be able to re-use some of it.


```python
def simulate_one_statistic():
    "Returns one value of our simulated test statistic"
    shuffled_labels = ...
    shuffled_tbl = ...
    group_means = ...
    ...
```

```python
grader.check("q3_9")
```

After you have defined your function, run the following cell a few times to see how the statistic varies.

```python
simulate_one_statistic()
```

### Question 3.10.

Complete the cell to simulate 4,000 values of the statistic. We have included the code that draws the empirical distribution of the statistic and shows the value of `observed_statistic_ab` from Question 3.6. Feel free to use as many lines of code as you need.

*Note:* This cell will take around a minute to run.


```python
repetitions = 4000

simulated_statistics_ab = make_array()
...
    simulated_statistics_ab = ...

# Do not change these lines
Table().with_columns('Simulated Statistic', simulated_statistics_ab).hist()
plt.scatter(observed_statistic_ab, -0.002, color='red', s=70);
```

```python
grader.check("q3_10")
```

### Question 3.11.

Use the simulation to find an empirical approximation to the p-value. Assign `p_val` to the appropriate p-value from this simulation. Then, assign `conclusion` to either `null_hyp` or `alt_hyp`. 

*Note:* Assume that we use the 5% cutoff for the p-value.


```python
# These are variables provided for you to use.
null_hyp = 'The data are consistent with the null hypothesis.'
alt_hyp = 'The data support the alternative more than the null.'

p_val = ...
conclusion = ...

p_val, conclusion # Do not change this line
```

```python
grader.check("q3_11")
```

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw06_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: If clicking on the link in this notebook results in an error, try downloading the file from the file explorer panel on the left.
5. Upload `hw06_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True)
```

Tests for q1_1 (points: [0, 0, 0, 0])
```
>>> isinstance(sample_size, (int, np.integer))
True
```
failure_message: ❌ `sample_size` is not an integer. It counts people, so it should be a whole number.
success_message: ✅ `sample_size` is a whole number
```
>>> isinstance(percent_V1, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ `percent_V1` is not a number. Compute it rather than storing it as a string.
success_message: ✅ `percent_V1` is a number
```
>>> 100 <= sample_size <= 500
True
```
failure_message: ❌ `sample_size` is outside the range we expect. It should be everyone vaccinated that week, both groups combined.
success_message: ✅ `sample_size` is in a plausible range for this clinic
```
>>> 0 <= percent_V1 <= 100
True
```
failure_message: ❌ `percent_V1` is not between 0 and 100. The note asks for a percent, so multiply the proportion by 100.
success_message: ✅ `percent_V1` is expressed as a percent between 0 and 100


Tests for q1_10 (points: [0, 0])
```
>>> isinstance(correct_doctor, (int, np.integer))
True
```
failure_message: ❌ `correct_doctor` is not an integer. Assign 1 or 2, not the text of the statement.
success_message: ✅ `correct_doctor` is an integer
```
>>> any((correct_doctor == x for x in (1, 2)))
True
```
failure_message: ❌ `correct_doctor` is not 1 or 2. Compare your p-value against the 5% cutoff, then pick the matching statement.
success_message: ✅ You picked one of the two statements


Tests for q1_5 (points: [0, 0])
```
>>> isinstance(observed_statistic, (float, np.floating))
True
```
failure_message: ❌ `observed_statistic` is not a number. It should be the value of the test statistic, not a description of it.
success_message: ✅ `observed_statistic` is a number
```
>>> 0 <= observed_statistic <= 100
True
```
failure_message: ❌ `observed_statistic` is outside the range we expect. It is a distance in percentage points, so it cannot be negative or above 100.
success_message: ✅ `observed_statistic` is on the percentage-point scale


Tests for q1_6 (points: [0, 0])
```
>>> isinstance(assumption_needed, (int, np.integer))
True
```
failure_message: ❌ `assumption_needed` is not an integer. Assign the number of the option you picked, not its text.
success_message: ✅ `assumption_needed` is an integer
```
>>> 1 <= assumption_needed <= 4
True
```
failure_message: ❌ `assumption_needed` is not 1, 2, 3, or 4. Pick the number of one of the four assumptions listed.
success_message: ✅ You picked one of the four options


Tests for q1_9 (points: [0, 0])
```
>>> import numpy as np
>>> isinstance(p_value, (float, np.floating))
True
```
failure_message: ❌ `p_value` is not a number. Compute it from your simulated statistics.
success_message: ✅ `p_value` is a number
```
>>> 0 <= p_value <= 1
True
```
failure_message: ❌ `p_value` is not between 0 and 1. Divide the count of simulated statistics by `num_simulations`.
success_message: ✅ `p_value` is a proportion between 0 and 1


Tests for q2_2 (points: None)
```
>>> import numpy as np
>>> isinstance(observed_tvd, (float, np.floating))
True
```
failure_message: ❌ `observed_tvd` is not a number. Make sure `calculate_tvd` returns a value rather than printing one.
success_message: ✅ `observed_tvd` is a number


Tests for q2_3 (points: None)
```
>>> len(simulated_tvds) == 10000
True
```
failure_message: ❌ `simulated_tvds` does not have 10,000 values. Are you appending each TVD inside the loop rather than overwriting the array?
success_message: ✅ `simulated_tvds` has 10,000 values
```
>>> all(simulated_tvds >= 0)
True
```
failure_message: ❌ At least one simulated TVD is negative. A total variation distance cannot be negative — check for a missing `abs`.
success_message: ✅ Every simulated TVD is non-negative
```
>>> len(np.unique(simulated_tvds)) != 1
True
```
failure_message: ❌ Every simulated TVD is identical. Are you re-sampling inside the loop?
success_message: ✅ Your simulated TVDs vary from trial to trial
```
>>> np.mean(simulated_tvds) < 0.1
True
```
failure_message: ❌ Your simulated TVDs are larger than expected under the null model. Check the sample size passed to `sample_proportions`.
success_message: ✅ Your simulated TVDs are in a plausible range


Tests for q2_4 (points: None)
```
>>> 0 <= p_value_tvd <= 1
True
```
failure_message: ❌ `p_value_tvd` is not between 0 and 1. A p-value is a proportion — divide by the number of simulations.
success_message: ✅ `p_value_tvd` is between 0 and 1


Tests for q3_1 (points: [0])
```
>>> 0 <= num_females <= 500
True
```
failure_message: ❌ `num_females` is outside the range of the sample. Remember `.group` sorts categories alphabetically, so "female" comes first.
success_message: ✅ `num_females` is a plausible count for this sample


Tests for q3_10 (points: None)
```
>>> len(simulated_statistics_ab) == 4000
True
```
failure_message: ❌ `simulated_statistics_ab` does not have 4,000 values. Check that you append inside the loop.
success_message: ✅ `simulated_statistics_ab` has 4,000 values


Tests for q3_11 (points: [0, 0, 0])
```
>>> import numpy as np
>>> isinstance(p_val, (float, np.floating))
True
```
failure_message: ❌ `p_val` is not a number. Compute it from `simulated_statistics_ab`.
success_message: ✅ `p_val` is a number
```
>>> 0 <= p_val <= 1
True
```
failure_message: ❌ `p_val` is not between 0 and 1. Divide the count of simulated statistics by the number of repetitions.
success_message: ✅ `p_val` is a proportion between 0 and 1
```
>>> type(conclusion) == str
True
```
failure_message: ❌ `conclusion` is not a string. Assign either `null_hyp` or `alt_hyp`, without quotes around the variable name.
success_message: ✅ `conclusion` is one of the two hypothesis strings


Tests for q3_2 (points: [0])
```
>>> import numpy as np
>>> isinstance(avg_male_vs_female, (bool, np.bool_))
True
```
failure_message: ❌ `avg_male_vs_female` is not True or False. Write a comparison between the two group averages rather than typing the answer.
success_message: ✅ `avg_male_vs_female` is a True/False value


Tests for q3_3 (points: [0, 0, 0, 0, 0])
```
>>> isinstance(null_statement_number, (int, np.integer))
True
```
failure_message: ❌ `null_statement_number` is not an integer. Assign the number of the statement, not its text.
success_message: ✅ `null_statement_number` is an integer
```
>>> isinstance(alternative_statement_number, (int, np.integer))
True
```
failure_message: ❌ `alternative_statement_number` is not an integer. Assign the number of the statement, not its text.
success_message: ✅ `alternative_statement_number` is an integer
```
>>> any((null_statement_number == x for x in np.arange(1, 7)))
True
```
failure_message: ❌ `null_statement_number` is not between 1 and 6. Pick the number of one of the statements listed.
success_message: ✅ `null_statement_number` is one of the six statements
```
>>> any((alternative_statement_number == x for x in np.arange(1, 7)))
True
```
failure_message: ❌ `alternative_statement_number` is not between 1 and 6. Pick the number of one of the statements listed.
success_message: ✅ `alternative_statement_number` is one of the six statements
```
>>> null_statement_number != alternative_statement_number
True
```
failure_message: ❌ You picked the same statement for both hypotheses. The null and the alternative have to be different claims.
success_message: ✅ You chose two different statements


Tests for q3_4 (points: [0, 0])
```
>>> isinstance(permutation_test_reason, (int, np.integer))
True
```
failure_message: ❌ `permutation_test_reason` is not an integer. Assign the number of the reason, not its text.
success_message: ✅ `permutation_test_reason` is an integer
```
>>> any((permutation_test_reason == x for x in np.arange(1, 4)))
True
```
failure_message: ❌ `permutation_test_reason` is not 1, 2, or 3. Pick the number of one of the reasons listed.
success_message: ✅ You picked one of the three reasons


Tests for q3_5 (points: [0, 0])
```
>>> isinstance(correct_test_stat, (int, np.integer))
True
```
failure_message: ❌ `correct_test_stat` is not an integer. Assign 1 or 2, not the text of the option.
success_message: ✅ `correct_test_stat` is an integer
```
>>> any((correct_test_stat == x for x in np.arange(1, 3)))
True
```
failure_message: ❌ `correct_test_stat` is not 1 or 2. Pick the statistic for which large values favor the alternative.
success_message: ✅ You picked one of the two test statistics


Tests for q3_6 (points: [0, 0])
```
>>> isinstance(observed_statistic_ab, (float, np.floating))
True
```
failure_message: ❌ `observed_statistic_ab` is not a number. Compute the statistic from the group averages.
success_message: ✅ `observed_statistic_ab` is a number
```
>>> observed_statistic_ab >= 0
True
```
failure_message: ❌ `observed_statistic_ab` is negative. Check the order of the subtraction — large values should favor the alternative you chose in Question 3.5.
success_message: ✅ `observed_statistic_ab` has the sign we expect for this sample


Tests for q3_7 (points: [0, 0])
```
>>> original_with_shuffled_labels.labels == ('Gender', 'Age', 'Shuffled Label')
True
```
failure_message: ❌ The table does not have the columns `Gender`, `Age` and `Shuffled Label`. Only assign `shuffled_labels`; the provided code builds the table.
success_message: ✅ The table has the three expected columns
```
>>> original_with_shuffled_labels.num_rows == 500
True
```
failure_message: ❌ The table does not have 500 rows. Shuffle the labels without replacement so each one is used exactly once.
success_message: ✅ Every row of the sample kept its place


Tests for q3_8 (points: [0, 0])
```
>>> isinstance(correct_q8, (int, np.integer))
True
```
failure_message: ❌ `correct_q8` is not an integer. Assign the number of the option, not its text.
success_message: ✅ `correct_q8` is an integer
```
>>> any((correct_q8 == x for x in np.arange(1, 4)))
True
```
failure_message: ❌ `correct_q8` is not 1, 2, or 3. Reason it out rather than computing it — pick the number of one of the options.
success_message: ✅ You picked one of the three options


Tests for q3_9 (points: None)
```
>>> -5 < simulate_one_statistic() < 5
True
```
failure_message: ❌ `simulate_one_statistic` returned a value far outside the range we expect. Make sure it returns a difference of group means.
success_message: ✅ `simulate_one_statistic` returns a plausible value
