```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 05: Probability, Simulation, Estimation, and Assessing Models

**Reading**: 
* [Randomness](https://inferentialthinking.com/chapters/09/Randomness.html) 
* [Sampling and Empirical Distributions](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Testing Hypotheses](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

<div style="background-color:#f2d8b2; border-left:6px solid #b36a00; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b> ℹ️ Appropriate collaboration and AI usage</b><br>
<span style="font-size:0.9em;">Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.</div></span>
</div>


<div style="background-color:#dac1d6; border-left:6px solid #853175; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>⚠️ This homework has hidden tests in it!</b><br>
<span style="font-size:0.9em;">Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**</span>
</div>

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Roulette


A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="roulette_wheel.jpeg" alt="Photograph of a roulette wheel showing the numbered pockets (0, 00, and 1 through 36) colored red, black, and green." width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

## Betting on Red
If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- If the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

### Question 1.1.

Define a function `dollar_bet_on_red` that takes the name of a color and returns your gain in dollars if that color had won and you had placed a one-dollar bet on red. Remember that the gain can be negative. Make sure your function returns an integer.

*Note:* You can assume that the only colors that will be passed as arguments are red, black, and green. Your function doesn't have to check that.


```python
def dollar_bet_on_red(...):
    ...
```

```python
grader.check("q1_1")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

### Question 1.2. 

Add a column labeled `Winnings: Red` to the table `wheel`. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. Your code should use the function `dollar_bet_on_red`.


```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("q1_2")
```

## Simulating 10 Bets on Red
Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

### Question 1.3.

Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. Your table should have the same three column labels as in `wheel`. Once you've created that table, set `sum_bets` to your net gain in all 10 bets, assuming that you bet one dollar on red each time.

*Hint:* It may be helpful to print out `ten_bets` after you create it!


```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("q1_3")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

### Question 1.4.

Let's see what would happen if you made more bets. Define a function `net_gain_red` that takes the number of bets and returns the net gain in that number of one-dollar bets on red.

*Hint:* You should use your `wheel` table within your function.


```python
def net_gain_red(...):
    ...
```

```python
grader.check("q1_4")
```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_red(10)
```

### Question 1.5.

Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

*Hint:* Think about which computational tool might be helpful for simulating a process multiple times. Lab 5 might be a good resource to look at!

*Note:* This cell might take a few seconds to run.


```python
num_bets = ...
repetitions = ...

all_gains_red = ...
...

len(all_gains_red) # Do not change this line! Check that all_gains_red is length 10000.
```

```python
grader.check("q1_5")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

### Question 1.6:

Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question.


```python
loss_more_than_50 = ...
```

```python
grader.check("q1_6")
```

## Betting on a Split
If betting on red doesn't seem like a good idea, maybe a gambler might want to try a different bet. A bet on a *split* is a bet on two consecutive numbers such as 5 and 6. This bet pays 17 to 1. That means if you place a one-dollar bet on the split 5 and 6, then:

- If the winning pocket is either 5 or 6, your gain is 17 dollars.
- If any other pocket wins, you lose your dollar, so your gain is -1 dollars.

### Question 1.7. 

Define a function `dollar_bet_on_split` that takes a pocket number and returns your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.

*Hint:* Remember that the pockets are represented as strings.


```python
def dollar_bet_on_split(...):
    ...
```

```python
grader.check("q1_7")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_split('5'))
print(dollar_bet_on_split('6'))
print(dollar_bet_on_split('00'))
print(dollar_bet_on_split('23'))
```

### Question 1.8.

Add a column `Winnings: Split` to the `wheel` table. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.


```python
split_winnings = ...
wheel = ...
wheel.show(5) # Do not change this line.
```

```python
grader.check("q1_8")
```

### Question 1.9.

Simulate the net gain in 200 one-dollar bets on the 5-6 split, repeating the process 10,000 times and saving your gains in the array `all_gains_split`.

*Hint:* Your code in Questions 1.4 and 1.5 may be helpful here!


```python
all_gains_split = ...

...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Split', all_gains_split)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("q1_9")
```

### Question 1.10.

Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram.

1. If you bet one dollar 200 times on a split, your chance of losing money is more than 50%.
2. If you bet one dollar 200 times in roulette, your chance of making more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.
3. If you bet one dollar 200 times in roulette, your chance of losing more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.

*Hint:* We've already seen one of these statements in a prior question.


```python
histogram_statements = ...
```

```python
grader.check("q1_10")
```

If this exercise has put you off playing roulette, it has done its job. If you are still curious about other bets, [the Wikipedia roulette bet-odds table](https://en.wikipedia.org/wiki/Roulette#Bet_odds_table) lists them all, and [the Wikipedia roulette house-edge section](https://en.wikipedia.org/wiki/Roulette#House_edge) is the bad news. The house – that is, the casino – always has an edge over the gambler.

## 2. Chances

Before you do this exercise, make sure you understand the logic behind all the examples in [Section 9.5](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen.
- Seeing if it is easier to find the chance that the event does not happen.

## Finding Chances

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green. In each part below, write an expression that evaluates to the chance of the event described.

### Question 2.1.

The winning pocket is black on all of the first three spins.


```python
first_three_black = ...
```

```python
grader.check("q2_1")
```

### Question 2.2.

The color green never wins in the first 10 spins.


```python
no_green = ...
```

```python
grader.check("q2_2")
```

### Question 2.3.

The color green wins at least once on the first 10 spins.


```python
at_least_one_green = ...
```

```python
grader.check("q2_3")
```

### Question 2.4.

Two of the three colors never win in the first 10 spins.

*Hint:* Imagine the event with the colors red, yellow, and orange. What situation(s) lead to two of the three colors never winning in the first 10 spins?


```python
lone_winners = ...
```

```python
grader.check("q2_4")
```

## Comparing Chances
In each of Questions 2.5–2.7, two events A and B are described. Choose from one of the following three options and set each answer variable to a single integer:

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time, and by thinking about the [law of averages](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html#the-law-of-averages).

### Question 2.5.

A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo.

- Event A: all four different animals are picked (assuming the child picks without replacement)
- Event B: all four different animals are picked (assuming the child picks with replacement)


```python
toys_option = ...
```

```python
grader.check("q2_5")
```

### Question 2.6.

In a lottery, two numbers are drawn at random with replacement from the integers 1 through 1000.

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


```python
lottery_option = ...
```

```python
grader.check("q2_6")
```

### Question 2.7.

A fair coin is tossed repeatedly.

- Event A: There are 60 or more heads in 100 tosses
- Event B: There are 600 or more heads in 1000 tosses

*Hint*: Think about the law of averages!


```python
coin_option = ...
```

```python
grader.check("q2_7")
```

## 3. Three Ways Python Draws Random Samples

You have learned three ways to draw random samples using Python:

- `tbl.sample` draws a random sample of rows from the table `tbl`. The output is a table consisting of the sampled rows. 

- `np.random.choice` draws a random sample from a population whose elements are in an array. The output is an array consisting of the sampled elements.

- `sample_proportions` draws from a categorical distribution whose proportions are in an array. The output is an array consisting of the sampled proportions in all the categories. 

```python
# Look through this code and run this cell for Questions 3.1 and 3.2
top = Table.read_table('top_movies_2017.csv').select(0, 1)
top.show(3)
```

```python
# Look through this code and run this cell for Questions 3.1 and 3.2
studios_with_counts = top.group('Studio').sort('count', descending=True)
studios_with_counts.show(3)
```

```python
# Look through this code and run this cell for Questions 3.1 and 3.2
studios_of_all_movies = top.column('Studio')
distinct_studios = studios_with_counts.column('Studio')

print("studios_of_all_movies:", studios_of_all_movies[:10], "...")
print("\n distinct_studios:", distinct_studios)
```

```python
# Look through this code and run this cell for Questions 3.1 and 3.2
studio_counts_only = studios_with_counts.column('count')
studio_proportions_only = studio_counts_only / sum(studio_counts_only)

print("studio_counts_only:", studio_counts_only)
print("\n studio_proportions_only:", studio_proportions_only)
```

In Questions 3.1 and 3.2, we will present a scenario. Determine which three of the following six options are true in regards to what the question is asking, and list them in the following answer cell. If your answer includes any of (i)-(iii), state what you would fill in the blank to make it true: `top`, `studios_with_counts`, `studios_of_all_movies`, `distinct_studios`, `studio_counts_only` or `studio_proportions_only`.

(i) This can be done using `sample` and the table _________.

(ii) This can be done using `np.random.choice` and the array ________.

(iii) This can be done using `sample_proportions` and the array _______.

(iv) This cannot be done using `sample` and the data given.

(v) This cannot be done using `np.random.choice` and the data given.

(vi) This cannot be done using `sample_proportions` and the data given.

<!-- BEGIN QUESTION -->

### Question 3.1.

Simulate a sample of 10 movies drawn at random with replacement from the 200 movies. Output True if Paramount appears more often than Warner Brothers among studios that released the sampled movies, and False otherwise.

*Example Answer:* (i) studios_of_all_movies, (iii) top, (v)

***Note***: Do not explain your answer for any of the options you've chosen; please follow the structure of the example answer provided.


<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3.2.

Simulate a sample of 10 movies drawn at random with replacement from the 200 movies. Output True if the first sampled movie was released by the same studio as the last sampled movie.

*Example Answer:* (i) studios_of_all_movies, (iii) top, (v)

***Note***: Do not explain your answer for any of the options you've chosen; please follow the structure of the example answer provided.


<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

## 4. Assessing Jade's Models

Before you begin, [Section 10.4](https://inferentialthinking.com/chapters/10/4/Random_Sampling_in_Python.html) of the textbook is a useful reference for this part.

## Games with Jade
Our friend Jade comes over and asks us to play a game with her. The game works like this: 

> We will draw randomly with replacement from a simplified 13 card deck with 4 face cards (A, J, Q, K), and 9 numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10). If we draw cards with replacement 13 times, and if the number of face cards is greater than or equal to 4, we lose.
> 
> Otherwise, Jade loses.

We play the game once and we lose, observing 8 total face cards. We are angry and accuse Jade of cheating! Jade is adamant, however, that the deck is fair.

Jade's model claims that there is an equal chance of getting any of the cards (A, 2, 3, 4, 5, 6, 7, 8, 9, 10, J, Q, K), but we do not believe her. We believe that the deck is clearly rigged, with face cards (A, J, Q, K) being more likely than the numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10).

### Question 4.1.

Assign `deck_model_probabilities` to a two-item array containing the chance of drawing a face card as the first element, and the chance of drawing a numbered card as the second element under Jade's model. Since we're working with probabilities, make sure your values are between 0 and 1.


```python
deck_model_probabilities = ...
deck_model_probabilities
```

```python
grader.check("q4_1")
```

### Question 4.2.

We believe Jade's model is incorrect. In particular, we believe there to be a *larger* chance of getting a face card. Which of the following statistics can we use during our simulation to test between the model and our alternative? Assign `statistic_choice` to the correct answer.

1. The distance (absolute value) between the actual number of face cards in 13 draws and 4, the expected number of face cards in 13 draws
2. The expected number of face cards in 13 draws, 4
3. The number of face cards we get in 13 draws


```python
statistic_choice = ...
statistic_choice
```

```python
grader.check("q4_2")
```

### Question 4.3.

Define the function `deck_simulation_and_statistic`, which, given a sample size and an array of model proportions (like the one you created in Question 4.1), returns the number of face cards in one simulation of drawing cards under the model specified in `model_proportions`.

*Hint:* Think about how you can use the function `sample_proportions`. 


```python
def deck_simulation_and_statistic(sample_size, model_proportions):
    ...

deck_simulation_and_statistic(13, deck_model_probabilities)
```

```python
grader.check("q4_3")
```

### Question 4.4.

Use your function from above to simulate the drawing of 13 cards 5000 times under the proportions that you specified in Question 4.1. Keep track of all of your statistics in `deck_statistics`.


```python
repetitions = 5000 
...

deck_statistics
```

```python
grader.check("q4_4")
```

Let’s take a look at the distribution of simulated statistics.

```python
# Draw a distribution of statistics 
Table().with_column('Deck Statistics', deck_statistics).hist()
```

<!-- BEGIN QUESTION -->

### Question 4.5.

Given your observed value, do you believe that Jade's model is reasonable, or is our alternative (that our deck is rigged) more likely? Explain your answer using the histogram produced above.


<div style="background-color:#fff3cd; border-left:6px solid #e0a800; color:#212529; padding:12px 16px; margin:8px 0; border-radius:4px;">
<b>✏️ YOUR TURN — replace this entire cell with your answer.</b><br>
<span style="font-size:0.9em;">Double-click to edit, delete this box, type your response, then Shift+Enter.</span>
</div>

<!-- END QUESTION -->

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw05_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: If clicking on the link in this notebook results in an error, try downloading the file from the file explorer panel on the left.
5. Upload `hw05_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True)
```

Tests for q1_1 (points: [0.5])
```
>>> dollar_bet_on_red('red') == 1
True
```
failure_message: ❌ `dollar_bet_on_red` does not return 1 when red wins. A winning one-dollar bet on red gains you a dollar.
success_message: ✅ `dollar_bet_on_red` pays out correctly when red wins


Tests for q1_10 (points: [0, 0])
```
>>> import numpy as np
>>> type(histogram_statements) == np.ndarray
True
```
failure_message: ❌ `histogram_statements` is not an array. Collect the statement numbers with `make_array(...)`.
success_message: ✅ `histogram_statements` is an array
```
>>> 1 <= histogram_statements.item(0) <= 3
True
```
failure_message: ❌ `histogram_statements` is empty or holds a number outside the range of the statements listed. Include at least one.
success_message: ✅ `histogram_statements` holds valid statement numbers


Tests for q1_2 (points: [0])
```
>>> wheel.labels[2] == 'Winnings: Red'
True
```
failure_message: ❌ `wheel` has no `Winnings: Red` column as its third column. Check the spelling, including the capital letters and the colon.
success_message: ✅ `wheel` has a `Winnings: Red` column in the right position


Tests for q1_3 (points: [0])
```
>>> set(['Pocket', 'Color', 'Winnings: Red']) == set(ten_bets.labels)
True
```
failure_message: ❌ Your sampled table does not have the same columns as `wheel`. If this is unexpected, restart the kernel and run the cells from the top.
success_message: ✅ Your simulated spins keep all three columns


Tests for q1_4 (points: [0])
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```
failure_message: ❌ `net_gain_red` returns the same value on every call. Sample the wheel inside the function so each call simulates fresh spins.
success_message: ✅ `net_gain_red` gives a different result each time, as a simulation should


Tests for q1_5 (points: [0])
```
>>> len(all_gains_red) == 10000
True
```
failure_message: ❌ `all_gains_red` does not hold 10,000 values. Append each simulated net gain inside the loop rather than overwriting the array.
success_message: ✅ `all_gains_red` holds 10,000 simulated net gains


Tests for q1_6 (points: [0])
```
>>> isinstance(loss_more_than_50, (bool, np.bool_))
True
```
failure_message: ❌ `loss_more_than_50` is not a boolean. Assign `True` or `False` directly, not the word in quotes.
success_message: ✅ `loss_more_than_50` is True or False


Tests for q1_7 (points: [0.5])
```
>>> dollar_bet_on_split('5') == 17
True
```
failure_message: ❌ `dollar_bet_on_split` does not return 17 when the 5 pocket wins. A winning split pays 17 to 1.
success_message: ✅ `dollar_bet_on_split` pays out correctly on a winning split


Tests for q1_8 (points: [0])
```
>>> wheel.labels[3] == 'Winnings: Split'
True
```
failure_message: ❌ `wheel` has no `Winnings: Split` column as its fourth column. Check the spelling, including the capital letters and the colon.
success_message: ✅ `wheel` has a `Winnings: Split` column in the right position


Tests for q1_9 (points: [0])
```
>>> len(all_gains_split) == 10000
True
```
failure_message: ❌ `all_gains_split` does not hold 10,000 values. Append each simulated net gain inside the loop.
success_message: ✅ `all_gains_split` holds 10,000 simulated net gains


Tests for q2_1 (points: [0])
```
>>> 0 < first_three_black < 1
True
```
failure_message: ❌ `first_three_black` is not between 0 and 1. Give the answer as a proportion, and remember all three spins have to be black.
success_message: ✅ `first_three_black` is a chance between 0 and 1


Tests for q2_2 (points: [0])
```
>>> 0 < no_green < 1
True
```
failure_message: ❌ `no_green` is not between 0 and 1. Give the answer as a proportion — the chance that none of the 10 spins is green.
success_message: ✅ `no_green` is a chance between 0 and 1


Tests for q2_3 (points: [0])
```
>>> 0 < at_least_one_green < 1
True
```
failure_message: ❌ `at_least_one_green` is not between 0 and 1. "At least one" is easiest to get from the chance of none.
success_message: ✅ `at_least_one_green` is a chance between 0 and 1


Tests for q2_4 (points: [0])
```
>>> 0 < lone_winners < 1
True
```
failure_message: ❌ `lone_winners` is not between 0 and 1. Count the ways exactly one color can win all 10 spins.
success_message: ✅ `lone_winners` is a chance between 0 and 1


Tests for q2_5 (points: [0, 0])
```
>>> isinstance(toys_option, (int, np.integer))
True
```
failure_message: ❌ `toys_option` is not an integer. Assign the number of the option, not its text.
success_message: ✅ `toys_option` is an integer
```
>>> 1 <= toys_option <= 3
True
```
failure_message: ❌ `toys_option` is not 1, 2, or 3. Compare the chances of events A and B and pick the matching statement.
success_message: ✅ You picked one of the three options


Tests for q2_6 (points: [0, 0])
```
>>> isinstance(lottery_option, (int, np.integer))
True
```
failure_message: ❌ `lottery_option` is not an integer. Assign the number of the option, not its text.
success_message: ✅ `lottery_option` is an integer
```
>>> 1 <= lottery_option <= 3
True
```
failure_message: ❌ `lottery_option` is not 1, 2, or 3. Compare the chances of events A and B and pick the matching statement.
success_message: ✅ You picked one of the three options


Tests for q2_7 (points: [0, 0])
```
>>> isinstance(coin_option, (int, np.integer))
True
```
failure_message: ❌ `coin_option` is not an integer. Assign the number of the option, not its text.
success_message: ✅ `coin_option` is an integer
```
>>> 1 <= coin_option <= 3
True
```
failure_message: ❌ `coin_option` is not 1, 2, or 3. Think about how the spread of the proportion of heads changes as the number of tosses grows.
success_message: ✅ You picked one of the three options


Tests for q4_1 (points: [0, 0])
```
>>> len(deck_model_probabilities) == 2
True
```
failure_message: ❌ `deck_model_probabilities` does not have exactly two entries — the chance of a face card and the chance of a numbered card.
success_message: ✅ `deck_model_probabilities` has two entries
```
>>> sum(deck_model_probabilities) == 1
True
```
failure_message: ❌ `deck_model_probabilities` does not sum to 1. Every card is either a face card or a numbered card.
success_message: ✅ `deck_model_probabilities` sums to 1, so it is a valid distribution


Tests for q4_2 (points: [0, 0])
```
>>> isinstance(statistic_choice, (int, np.integer))
True
```
failure_message: ❌ `statistic_choice` is not an integer. Assign the number of the statistic, not its text.
success_message: ✅ `statistic_choice` is an integer
```
>>> 1 <= statistic_choice <= 3
True
```
failure_message: ❌ `statistic_choice` is not 1, 2, or 3. Pick the statistic that gets larger when the deck favors face cards.
success_message: ✅ You picked one of the three statistics


Tests for q4_3 (points: [0])
```
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> 0 <= num_face <= 13
True
```
failure_message: ❌ `deck_simulation_and_statistic` returned a value outside 0 to 13. With a sample size of 13 you cannot draw more face cards than that.
success_message: ✅ `deck_simulation_and_statistic` returns a plausible number of face cards


Tests for q4_4 (points: [0, 0])
```
>>> len(deck_statistics) == repetitions
True
```
failure_message: ❌ `deck_statistics` does not have `repetitions` values. Append each simulated statistic inside the loop.
success_message: ✅ `deck_statistics` has one value per repetition
```
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```
failure_message: ❌ At least one value in `deck_statistics` is outside 0 to 13. Check the sample size you pass to your function.
success_message: ✅ Every simulated statistic is a plausible face-card count
